---
title: "QSS Tutorial 1: Causality I"
output:
  learnr::tutorial:
    progressive: true
    ace_theme: github
    theme: united
runtime: shiny_prerendered
description: >
  This tutorial will cover logicals, subsetting, factors, and
  calculating means in experiments. Companion to Chapter 2 of
  "Quantitative Social Science: An Introduction".
---

```{r setup, include=FALSE}
library(gradethis)
library(learnr)
library(qsslearnr)
tutorial_options(exercise.checker = gradethis::grade_learnr)
knitr::opts_chunk$set(echo = FALSE)
tut_reptitle <- "QSS Tutorial 1: Output Report"
data(resume, package = "qss")
```



## Conceptual Questions


```{r concept-quiz, echo=FALSE}
quiz(
  caption = "",
  question(
    "Suppose a variable is binary, that is, it takes on values of either 0 or 1 (for example, female gender). Which of the following is the same as its sample mean?",
    answer("the sample median"),
    answer("the sample proportion of 1s", correct = TRUE),
    answer("neither of these")
  ),
  question(
    "What kind of value is `FALSE`?",
    answer("character"),
    answer("logical", correct = TRUE),
    answer("binary"),
    answer("numeric")
  ),
  question(
    "In order to calculate the mean of a variable we have used the `length()` function in the denominator. The `length()` of a vector is equivalent to:",
    answer("the number of elements", correct = TRUE),
    answer("the height"),
    answer("the maximum")
  ),
  question(
    "How are factor variables different from categorical variables?",
    answer("They are the same", correct = TRUE),
    answer("Factor variables contain numeric values"),
    answer("Categorical variables tend to have more levels or categories")
  )
)
```


## Working with Logicals in R

### Exploring the resume data


In this tutorial, we are going to be working with the resume data from Section 2.1 of QSS. This data comes from an experiment where researchers sent fictitious resumes with different names that implied different race and gender combinations to see if potential employers were more likely to call back names associated with different racial groups and genders. 

Let's first explore the data a bit. It's stored as `resume`.

#### Exercise 

- Use the `head` function to show the first six lines of the `resume` data. 


```{r head, exercise=TRUE}
## print the first 6 lines of the data

```

```{r head-solution}
head(resume)
```

```{r head-check}
grade_code()
```

- Print out the dimension of the `resume` data. 


```{r resume-dim, exercise=TRUE}

```

```{r resume-dim-solution}
dim(resume)
```

```{r resume-dim-check}
grade_code()
```


- Use the `summary` function to show a summary of the data.


```{r res-summary, exercise=TRUE}

```

```{r res-summary-solution}
summary(resume)
```

```{r res-summary-check}
grade_code()
```


### Creating a cross tab


To help you analyze this data, you can use a cross tabulation. Cross tabulation (or contingency table) is a table that quickly summarizes categorical data. For instance, in the resume data, we have a `sex` variable that tells us whether or not the fictitious resume had a male or a female name.

#### Exercise

- Using the `table` function, create a cross tab of the `sex` and `call` variables in the resume data.

```{r table, exercise = TRUE}

```

```{r table-check}
grade_result(
  pass_if(~ identical(.result, table(resume$sex, resume$call))),
  pass_if(~ identical(.result, table(resume$call, resume$sex)))
)
```

### Logical values


Pretty soon, you'll be doing more complicated subsetting in R. To do this, it's helpful to understand a special type of object in R: the **logical**. There are two values associated with this type of object: `TRUE` and `FALSE` (where the uppercase is very important).

#### Exercises

- Create a vector called `x` that contains two `TRUE` values and two `FALSE` values in that order. 
- Find the sum of that vector. 

```{r logical, exercise = TRUE}
## creat a vector with two TRUE values and two FALSE values
x <- 

## take the sum of this vector

```


```{r logical-hint}
x <- c()

mean(x)
```


```{r logical-check}
grade_result_strict(
  pass_if(~ identical(x, c(TRUE, TRUE, FALSE, FALSE))),
  pass_if(~ identical(.result, 2L))
)
```


- Create a vector called `z` that contains one `TRUE` values and three `FALSE` values in that order. 
- Find the mean of that vector. 

```{r logical2, exercise = TRUE}
## creat a vector with one TRUE values and three FALSE values
z <- 

## take the mean of this vector

```

```{r logical2-hint}
z <- c()

mean(z)
```


```{r logical2-check}
grade_result_strict(
  pass_if(~ identical(z, c(TRUE, FALSE, FALSE, FALSE))),
  pass_if(~ identical(.result, 0.25))
)
```


### Comparing logicals


We often combine logical statements using AND (`&`) and OR (`|`) in R. For AND statements, both expressions have to be true for the whole expression to be true:

- `TRUE & FALSE`, `FALSE & TRUE`, and `FALSE & FALSE` are `FALSE`
- `TRUE & TRUE` is `TRUE`

For OR statements, either statement being true makes the whole expression true:

- `TRUE | FALSE`, `FALSE | TRUE`, and `TRUE | TRUE` are `TRUE`
- `FALSE | FALSE` is `FALSE`


```{r comparing-logicals, echo=FALSE}
question("What does expression `(TRUE | FALSE) & TRUE` evaluate to?",
  answer("`TRUE`", correct = TRUE),
  answer("`FALSE`"),
  answer("`NA`")
)
```


### Comparing objects


There are several *relational operators* that allow us to compare objects in R. The most useful of these are the following:

- `>` greater than, `>=` greater than or equal to
- `<` less than, `<=` less than or equal to
- `==` equal to
- `!=` not equal to

When we use these to compare two objects in R, we end us with a logical object. You can also compare a vector to a particular number.


#### Exercises

- Test if 10 is greater than 5. 

```{r gt-test, exercise=TRUE}

```

```{r gt-test-solution}
10 > 5
```

```{r gt-test-check}
grade_code()
```

- Test which values in the vector `x` is greater than or equal to 0.

```{r compare, exercise = TRUE}
## x vector
x <- c(-2, -1, 0, 1, 2)

## test which values of x are greater than or equal to 0

```

```{r compare-check}
grade_result(
  fail_if(~ identical(.result, x > 0),
          "Did you forget the 'or equal to' part of the comparison?"),
  pass_if(~ identical(.result, x >= 0))
)
```

## Subsets in R

### Subsetting a data frame


You can use the same logical statements you have been using to create subsets of a data frame. These can often be helpful because we'll want to calculate various quantities of interest for different subsets of the data. For this exercise, we will use the `resume` data frame made up of the variables `firstname`, `sex`, `race`, and `call`. As a reminder, here is what the data look like:

```{r}
resume
```

#### Exercise 

- Use the `subset` function to create a subset of the `resume` data frame that is only female names that sound white. Save this subset as `resume.wf`
- Use the `head` function to print out the first 6 lines of this subset. 
- Calculate the mean of the `call` variable in this subset.


```{r subset, exercise = TRUE}
## create the subset for white female names and
## assign it to resume.wf
resume.wf <- ...

## print the first 6 lines of the subset


## calculate the mean of the callback variable (call)

```

```{r subset-solution}
resume.wf <- subset(resume, subset = (race == "white" & sex == "female"))
head(resume.wf)
mean(resume.wf$call)
```

```{r subset-check}
grade_result(
  pass_if(~ identical(.result, mean(subset(resume, subset = (race == "white" & sex == "female"))$call)))
)
```


### Comparing means across treatment conditions


You can use the same ideas as in the last step to create a different subset of the data corresponding to white-sounding female names. Then, you can compare the average callback for the white-female names to the average callback for the black-female names. This will give you a sense of how the employer callback rate varies by racial group of the applicant for females.

#### Exercise

- Create a subset of the `resume` data for black-sounding female names. 
- Print the difference in means between the `call` variable in the white-sounding name subset and the black-sounding name subset.

```{r effect, exercise = TRUE}
## create the subset for white female names
resume.wf <- subset(resume, subset = (race == "white" & sex == "female"))

## create the subset for black female names
resume.bf <- 

## calculate the difference in callback means

```

```{r effect-solution}
## create the subset for white female names
resume.wf <- subset(resume, subset = (race == "white" & sex == "female"))

## create the subset for black female names
resume.bf <- subset(resume, subset = (race == "black" & sex == "female"))

## compare the difference in means
mean(resume.wf$call) - mean(resume.bf$call)
```

```{r effect-check}
grade_code("You just analyzed an experiment! Way to go!")
```


## Submit

```{r context="setup"}
submission_ui
```

```{r context="server"}
submission_server()
```
